<a href="https://colab.research.google.com/github/MChristian00/MChristian00/blob/main/assignment1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Christian Manzi"
COLLABORATORS = ""
```

---

<a href='https://ai.meng.duke.edu'> = <img align="left" style="padding-top:10px;" src=https://storage.googleapis.com/aipi_datasets/Duke-AIPI-Logo.png>

# Assignment 1
Welcome to your first assignment in Module 2 of the Summer Computing Institute!  This assignment is all about getting comfortable working in Jupyter Notebooks, an extremely popular medium for doing data science programming. As you read this you are seeing the benefits of working in notebooks - they allow you to mix text, code, and images all together in a single file.

Before you start on the assignment, be sure you have fully read through the "assignment_instructions.md" document and understand how to submit your work for grading.  

The questions below will ask you to write some code, and then in the cell directly following it there will be one or more test cases to evaluate your solution.  Some of these test cases will be visible to you so that you can verify your work, but some cells will also have additional hidden test cases which are applied when grading to keep everyone honest. 

Before you push your final version back to your GitHub repo, always make sure that:  
- You restart the kernel of your notebook (under the "Kernel" menu at top) and run everything from top to bottom to ensure you do not have any errors  
- Your code is passing all the visible test cases  

After you push your solution back to your GitHub repo, view the repo in your web browser to make sure that it contains your intended version.

### Question 1
Let's start with a warmup exercise to practice our Python.  We are going to perform some basic analysis of the stock price for a fictional company NewCo.  Let's suppose we receive some data on the last 10 days of NewCo's stock price.  The data is in the form of a dictionary 'newco_dict' where the keys are ['DATE','LOW','HIGH','CLOSE'].  'DATE' contains a list of string dates. 'LOW' and 'HIGH' each contain a list of the lowest price and highest price respectively of the stock during each day in the 10-day period. 'CLOSE' contains a list of the closing price of the stock at the end of each day during the 10-day period.

For example, newco_dict might look like the following:  

```
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}
```

Our goal in this assignment will be to find all the dates during the 10-day period when the closing price of the stock was equal to the mode (most frequent) value of the closing prices over the 10-day period.  We are going to tackle this in two parts.  First we will create a function to find the mode of the closing prices.  In part 2 we will then create a function to identify the dates on which the closing price was equal to the mode.


NOTE: For all parts of Question 1, you may not import any additional libraries - I would like you to solve these using only the Python you have learned so far and no helper libraries unless already imported for you. 

#### Question 1.1
Let's first define a function `most_freq_closing(price_dict)` which takes as input a dictionary of prices in the form of 'newco_dict' and returns the most frequent closing price (the mode of the closing prices) during the period as an integer.  You may assume that the dictionary is non-empty and that there is a single mode price (in real life, some datasets may have no mode, if all elements are represented equally as often, or may have more than one mode).

There are many ways to approach this problem.  One method (although there are others) might be to iterate through the set of the unique closing values and count how often each one appears during the 10 days.  The mode is the one that appears most frequently (has the highest count of appearances).  

How you structure your solution is up to you, although you must complete the `most_freq_closing` function below.  You may include all your solution code within that function or break sub-tasks out into separate smaller functions (which you can insert into the same cell or add a new cell).

<!-- grade_id: cell-c61d21d175f8cc1a, points: 0, grade: false, solution: true -->
```python
def most_freq_closing(price_dict):
    # Identifies the mode of the stock closing prices
    # YOUR CODE HERE
    modeIndex = None
    count = 0
    closeList = price_dict['CLOSE']
    for item in closeList:
        modeCount = closeList.count(item)
        if modeCount >= count:
            modeIndex = closeList.index(item)
            count = modeCount
            pass
        pass
    mode = closeList[modeIndex]
    return mode
    raise NotImplementedError()
```

<!-- grade_id: cell-1e5d837902352889, points: 1, grade: true, solution: false -->
```python
# `most_freq_closing()`: test cell
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}
mode_price = most_freq_closing(newco_dict)
print(f'Your function calculated the most frequent price as {mode_price}')
assert type(mode_price) == int, 'Your function did not return an integer'
assert mode_price == 75,'Incorrect, check your function and try again'
print('\n(Passed visible tests!)')

```

Output:
```
Your function calculated the most frequent price as 75

(Passed visible tests!)
```

#### Question 1.2
Let's now make use of our new function `most_freq_closing` to find all dates on which the most frequent closing price occurred.  Write a new function `dates_most_freq_closing(price_dict)` which finds and returns the dates on which the most frequent (mode) closing price occurs.  The input to the function is a dictionary of the form of `newco_dict` and the output is a **list** of dates.  Your function should make use of the function `most_freq_closing()` that you just wrote to find the mode closing price, and then you should find all dates on which it occurs.

<!-- grade_id: cell-2f9bc10ff3889fa4, points: 0, grade: false, solution: true -->
```python
def dates_most_freq_closing(price_dict):
    # Returns a list of dates on which the stock closing price was equal to the mode
    mode = most_freq_closing(price_dict)
    closeDateList = []
    dateList = price_dict['DATE']
    closeList = price_dict['CLOSE']

    for date in dateList:
        dateIndex = dateList.index(date)
        if closeList[dateIndex] == mode:
            closeDateList.append(date)
            pass
        pass
    return closeDateList
    # raise NotImplementedError()
```

<!-- grade_id: cell-2dddb819d9fe493a, points: 1, grade: true, solution: false -->
```python
# `dates_most_freq_closing()`: test cell
newco_dict = {'DATE':['06-14-21','06-15-21','06-16-21','06-17-21','06-18-21',
                      '06-21-21','06-22-21','06-23-21','06-24-21','06-25-21'],
              'LOW':[80,72,70,74,82,88,94,88,73,67],
              'HIGH':[95,81,76,88,90,96,102,94,90,72],
              'CLOSE':[90,75,75,86,90,93,100,91,75,68]}

# Check that dates_most_freq_closing uses most_freq_closing
orig_most_freq_closing = most_freq_closing
del most_freq_closing
try:
    dates_most_freq_closing(newco_dict)
except NameError:
    pass
else:
    raise AssertionError("dates_most_freq_closing does not use most_freq_closing")
finally:
    most_freq_closing = orig_most_freq_closing
    
# Check output
dates = dates_most_freq_closing(newco_dict)
print(f'Your function returned the dates: {dates}')
assert type(dates) == list, 'Your function did not return a list of dates'
assert type(dates[0] == str), 'The values in your list are not strings'
assert dates == ['06-15-21', '06-16-21', '06-24-21'], 'Incorrect, check your function and try again'
print('\n(Passed visible tests!)')
```

Output:
```
Your function returned the dates: ['06-15-21', '06-16-21', '06-24-21']

(Passed visible tests!)
```